```python
version = "v2.2.033020"
```

# Assignment 2: Mining Itemsets (Part V)


## Itemset Similarity
Recall from Week 1 that pattern and similarity are two basic outputs of data mining. So far, we have been playing with patterns - frequent itemsets and association rules can all be seen as "patterns". In the last part, let's work on itemset similarities.

### Exercise 6. Jaccard Similarity (20 pts)
Jaccard similarity is a simple but powerful measurement of itemset similarity, defined as follows:
$$\text{Jaccard_similarity(A, B)} = \frac{|A\cap B|}{|A\cup B|}$$

#### Exercise 6(a) (10 pts)
Complete the following function to calculate the Jaccard similarity between two sets, you may assume that at least one of the sets is not empty.

<!-- grade_id: cell-0b0adf575adc8978, points: 0, grade: false, solution: true -->
```python
def jaccard_similarity(set_a, set_b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9e152786ac353a99, points: 10, grade: true, solution: false -->
```python
# This code block tests whether the `jaccard_similarity` function work as expected.
# We hide some tests, so passing all the displayed assertions does not gurantee the bonus points.

assert abs(jaccard_similarity(set(['🍇', '🍈', '🍉']), set(['🍊', '🍋', '🍌'])) - 0) < 1e-8
assert abs(jaccard_similarity(set(['🍇', '🍈', '🍎']), set(['🍊', '🍋', '🍎'])) - 0.2) < 1e-8
assert abs(jaccard_similarity(set(['🍇', '🍒', '🍎']), set(['🍊', '🍒', '🍎'])) - 0.5) < 1e-8
assert abs(jaccard_similarity(set(['🍓', '🍒', '🍎']), set(['🍒', '🍎', '🍓'])) - 1) < 1e-8

```

#### Exercise 6(b) (10 pts)
A few questions to help you better understand the properties of Jaccard similarity:
1. What is the range of Jaccard similarity?
2. When is the max/min Jaccard similarity achieved?

Complete the following function to indicate the minimum and maximum value a Jaccard similarity score can achieve. (5 pts each).

<!-- grade_id: cell-6011d49670d5abec, points: 0, grade: false, solution: true -->
```python
# Please change min_value and max_value to your answers

min_value = -100
max_value = -100

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed7b0719031d74b1, points: 5, grade: true, solution: false -->
```python
# This code block evaluate if min_value is assigned with the correct value
```

<!-- grade_id: cell-d177e6fd17a56e22, points: 5, grade: true, solution: false -->
```python
# This code block evaluate if max_value is assigned with the correct value
```

With the Jaccard similarity calculated, we can now calculate the Jaccard similarity between any given Tweet with all other Tweets and find the Tweets that are most similar in terms of the set of food/drink emojis used. (Not graded.)

```python
import pandas as pd
import numpy as np

tweets_df = pd.read_csv("assets/food_drink_emoji_tweets.txt", sep="\t", header=None)
tweets_df.columns = ['text']

emoji_list = "🍇🍈🍉🍊🍋🍌🍍🥭🍎🍏🍐🍑🍒🍓🥝🍅🥥🥑🍆🥔🥕🌽🌶🥒🥬🥦🍄🥜🌰🍞🥐🥖🥨🥯🥞🧀🍖🍗🥩🥓🍔🍟🍕🌭🥪🌮🌯🥙🥚🍳🥘🍲🥣🥗🍿🧂🥫🍱🍘🍙🍚🍛🍜🍝🍠🍢🍣🍤🍥🥮🍡🥟🥠🥡🦀🦞🦐🦑🍦🍧🍨🍩🍪🎂🍰🧁🥧🍫🍬🍭🍮🍯🍼🥛☕🍵🍶🍾🍷🍸🍹🍺🍻🥂🥃"
emoji_set = set(emoji_list)

def extract_uniq_emojis(text):
    return 

tweets_df['emojis'] = tweets_df.text.apply(lambda text:np.unique([chr for chr in text if chr in emoji_set]))

tweets_df['jaccard'] = tweets_df.emojis.apply(lambda x:jaccard_similarity(set(tweets_df.loc[0].emojis), set(x)))
tweets_df.sort_values('jaccard',ascending=False).head(n=10)
```

**🍾🍾 Congratulations 🎉🎉, you have now finished up all the exercises in this assignment!**