```python
version = "v2.2.033020"
```

# Assignment 3: Mining Vectors and Matrices (Part III)

In Part II of the assignment, we have implemented a few similarity/distance functions. In part III, we will see how can we use them to analyze our restaurant rating data.

Let's first import the necessary packages and dependencies and load the dataset prepared in Part I of this assignment.

```python
import pandas as pd
import numpy as np

business_df = pd.read_csv('assets/montreal_business.csv')
business_df.set_index('business_id', inplace=True)

review_df = pd.read_csv('assets/montreal_user.csv')
rating_df = review_df.pivot_table(index=['business_id'], columns=['user_id'], values='stars')
rating_df.fillna(0, inplace=True)

missing_business_id = set(business_df.index) - set(rating_df.index)
business_df.drop(missing_business_id, inplace=True)

business_df.head()
```

```python
rating_df.head(5)
```

Suppose one of my favorite restaurants in Montréal is named *Modavie*, and I want to know how similar (in terms of customer ratings) it is compared to other restaurants. Let's see how we can do it.

```python
modavie_id = business_df[business_df.name.str.contains("Modavie")].index[0]
print(modavie_id)
```

This is a sanity check that restaurant "Modavie" does exist in our dataset. 

### Exercise 3. Find similar restaurants (15 pts)
Once we have a similarity metric implemented, we can use it to find the most similar vectors. In this case, can you find out which restaurants are most similar to Modavie based on dot product? 

More specifically, can you implement the `find_max_dot_prod_restaurants` function to return the business_id of five restaurants that have the **largest** dot product with Modavie?

To do this, you need to compute the dot product between every restaurant and Modavie, and store the dot products as a seperate column named `modavie_dot_prod` on the `business_df` dataframe. Then, you should output the five rows with the largest dot product (ranked in decreasing order and excluding Modavie itself). 

**HINT 1** You may refer to Assignment 2(Part V) Jaccard Similarity to think about how to calculate the similarity scores for all restaurants.

**HINT 2** Use `np.dot` (instead of implementing your own dot product) to **greatly** speed up the execution.

<!-- grade_id: cell-b54cefee95496e66, points: 0, grade: false, solution: true -->
```python
def find_max_dot_prod_restaurants(top_n):
    modavie_vector = rating_df.loc[modavie_id]
    # YOUR CODE HERE
    raise NotImplementedError()
```

We have provided the correct answers for you to verify your solution, if your code takes more than a few seconds to execute, please double check to ensure you are using `np.dot` instead of our DIY-ed `dot_prod`.

```python
max_sim_restaurants = find_max_dot_prod_restaurants(10)
max_sim_restaurants.name
```

<!-- grade_id: cell-f72957ef426e186d, points: 15, grade: true, solution: false -->
```python
answer = find_max_dot_prod_restaurants(5)
assert answer.iloc[0]['name'] == "Schwartz's"
assert answer.iloc[1]['name'] == "La Banquise"
assert answer.iloc[2]['name'] == "Olive & Gourmando"
assert answer.iloc[3]['name'] == "Maison Christian Faure"
assert answer.iloc[4]['name'] == "Reuben's Deli & Steakhouse"
```

### Exercise 4. (10 pts)
We can also characterize the similarity on a larger scale. In fact, we are curious to know if Modavie's good ratings are more similar to other local area restaurants or to other French restaurants.

In this exercise, please calculate the average **cosine similarity** between Modavie and (1) other restaurants in the same local area (defined by the **first 3 digits** of the postal code), (2) other French restaurants (restaurants with "French" in its `categories`)

**Note**:
1. Canada postal code uses a "AXB YCZ" pattern, where A, B, and C represent letters and X, Y, and Z represent numbers. For this exercise, we will only use the first 3 digits of a zip code, that is, "H2Y" for Modavie.
2.Again, you can see the wildness of real-world data: some food trucks are also listed as restaurants but they don't have fixed locations, let alone postal codes. For this exercise, we assume that they do not belong to any neighborhood.
3. Modavie itself should be excluded from the calculation.
4. We have prepared a NumPy-based `cosine_similarity` function for your convenience.

```python
business_df.loc[modavie_id].postal_code
```

<!-- grade_id: cell-5abc221ff5074456, points: 0, grade: false, solution: true -->
```python
def cosine_similarity(vec_x, vec_y):
    return np.dot(vec_x, vec_y)/(np.linalg.norm(vec_x) * np.linalg.norm(vec_y))
def similarity_with_local_restaurant():
    modavie_vector = rating_df.loc[modavie_id]
    # YOUR CODE HERE
    raise NotImplementedError()
    return avg_cos_sim_local, avg_cos_sim_french
```

```python
avg_cos_sim_local, avg_cos_sim_french = similarity_with_local_restaurant()
print(avg_cos_sim_local, avg_cos_sim_french)
```

<!-- grade_id: cell-2a7221a3b0e053d1, points: 5, grade: true, solution: false -->
```python
avg_cos_sim_local, avg_cos_sim_french = similarity_with_local_restaurant()
assert abs(avg_cos_sim_local - 0.019919428868434952) < 1e-8, "[Exercise 4] Wrong value for avg_cos_sim_local."

```

<!-- grade_id: cell-8f69f0d924b0057b, points: 5, grade: true, solution: false -->
```python
avg_cos_sim_local, avg_cos_sim_french = similarity_with_local_restaurant()
assert abs(avg_cos_sim_french - 0.01361887560334431) < 1e-8, "[Exercise 4] Wrong value for avg_cos_sim_french."
```

After this exercise, you should be able to compute the similarity between many data vectors and find the most similar vectors of a given vector. This can be used as the foundation of many advanced algorithms such as k-nearest neighbor classification, information retrieval, or clustering,  which are the core of search engines or recommender systems. Computing the average similarity to selected groups of vectors (e.g., French restaurants) also provides a powerful tool to get deeper understanding of the data and to generate new features for downstream machine learning tasks. We encourage you to try these similarity metrics on your own data sets. 