```python
version = "v2.2.033020"
```

# Assignment 3: Mining Vectors and Matrices (Part IV)

## Patterns in Matrix Data
In the last part of the assignment, let's explore the patterns in a matrix as a whole instead of in individual vectors. As discussed in class, one type of such patterns are *eigenvectors*, which can be obtained through *Singular Value Decomposition* (SVD):

$$X=U\Sigma V^T.$$

$X$ is an $n \times p$ matrix, $U\cdot U^T = I$, $V\cdot V^T = I$, and $\Sigma$ is an $n\times p$ diagonal matrix with non-zero singular values.

Let's first walk through the example in the lecture.

```python
import pandas as pd
import numpy as np
from sklearn.decomposition import TruncatedSVD

np.set_printoptions(precision=6)
np.set_printoptions(suppress=True)
```

```python
doc_word_df = pd.DataFrame([[1, 1, 1, 0, 0], 
                            [2, 2, 2, 0, 0],
                            [1, 1, 1, 0, 0],
                            [5, 5, 5, 0, 0],
                            [0, 0, 0, 2, 2],
                            [0, 0, 0, 3, 3],
                            [0, 0, 0, 1, 1]])
doc_word_df.columns = ['data', 'information', 'retrieval', 'brain', 'lung']
doc_word_df.index = ['Data Science 1', 'Data Science 2', 'Data Science 3, ', 'Data Science 4', 
                     'Medical Report 1', 'Medical Report 1', 'Medical Report 1']
doc_word_df
```

Several Python packages, such as NumPy, SciPy, and scikit-learn, all provide APIs for SVD. In this assignment, we will use `TruncatedSVD` API from scikit-learn, as it allows us to compute only the largest $k$ singular values in $\Sigma$ and the corresponding $k$ columns (rows) of $U$ ($V^T$), which is more efficient and practical in real-world applications. You can learn more about the API from its [documentation](https://scikit-learn.org/stable/modules/generated/sklearn.decomposition.TruncatedSVD.html).

For this example, we explicitly specify $k=2$ (so we are only interested in the first 2 components), and we assign a constant to `random_state` to ensure that the result is reproducible. 

```python
doc_word_svd = TruncatedSVD(n_components=2, random_state=0)
doc_word_svd.fit(doc_word_df)
```

Now the object `doc_word_svd` stores all the necessary results of the decomposed matrix. You can view the singular values (the diagonal values in $\Sigma$) in the `singular_values_` attribute and the row vectors of $V^T$ in the `components_` attribute.

```python
sigma_diag = doc_word_svd.singular_values_
v_t = doc_word_svd.components_
print(sigma_diag)
print(v_t)
```

You may wonder where the $U$ matrix is. Although $U$ is not explicitly stored in the `TruncatedSVD` object. it can be recovered with $X$, $\Sigma$, and $V$, based on the following formula:

$$ X^* = U\Sigma = XV.$$

Thus, $U$ can be written as

$$ U = X^* \Sigma^{-1} = XV\Sigma^{-1}. $$

We can verify this with the following code:

```python
x = doc_word_df.values
v = v_t.T
sigma_inverse = np.linalg.inv(np.diag(sigma_diag))

u = x.dot(v).dot(sigma_inverse)
print(u)
```

Now let's try to interpret $U$, $\Sigma$, and $V^T$ by thinking about the following questions (not graded):

* Which codes the Data Science / Medical concept of each document?
* Which codes the strength of each concept?
* Which codes the word representation of each concept?

One important application of SVD is to transform the original matrix $X$ in to a new matrix $X^*$. This transformation projects the $p$-dimensional row vectors in the original matrix into $k$-dimensional vectors in the new vector space. $X^*$ encodes most of the information in the original $X$ matrix with fewer dimensions.

Mathamatically, the transformation is made by post-multiplying $X$ with the $V$ matrix ($ X^* = XV$). The `TruncatedSVD` API has implemented such transformation as a `transform` function. You may verify this with the following code:

```python
print(doc_word_svd.transform(doc_word_df))
```

```python
print(doc_word_df.values.dot(v))
```

Now that you are familiar with the SVD operations. Let's apply it to the Montréal restaurant dataset. Run the following code block to load the dataset prepared in Part I of this assignment.

```python
business_df = pd.read_csv('assets/montreal_business.csv')
business_df.set_index('business_id', inplace=True)

review_df = pd.read_csv('assets/montreal_user.csv')
rating_df = review_df.pivot_table(index=['business_id'], columns=['user_id'], values='stars')
rating_df.fillna(0, inplace=True)

missing_business_id = set(business_df.index) - set(rating_df.index)
business_df.drop(missing_business_id, inplace=True)
```

### Exercise 5. (15 pts.)
Complete the following `transformed_rating` function, which transforms the original $n\times p$ rating matrix into a new $n \times k$ matrix using the `TruncatedSVD` API. The function should return the $n \times k$ matrix.

**Note:**
1. Please set the `random_state` to 0 so that the results do not change over different runs.
2. $k$ can take values between 1 and $p$, not necessarily the 2 used in the document-word matrix example.
3. You may either use the `transformation` function in the `TruncatedSVD` object or use the dot product. Both are OK.

<!-- grade_id: cell-6da0da7f356ffe5d, points: 0, grade: false, solution: true -->
```python
def svd_transformed_rating(original_matrix, k, random_state=0):
    # YOUR CODE HERE
    raise NotImplementedError()
    return transformed_matrix
```

With the `svd_transformed_rating` function, we can calculate the transformed matrix through the following command.

```python
rating_matrix_d2 = svd_transformed_rating(rating_df, 2)
# You can uncomment the following line to view the result
# rating_matrix_d2
```

<!-- grade_id: cell-abfaa1e6400f27cc, points: 15, grade: true, solution: false -->
```python
# This code block tests if the `svd_transformed_rating` function is implemented correctly.
# We hide some tests, so passing all the displayed assertions does not guarantee full points.

rating_matrix_d2 = svd_transformed_rating(rating_df, 2)
# d2_groundtruth is calculated via the following command:
# d2_groundtruth = svd_transformed_rating(rating_df, 4)[:4]
d2_groundtruth = np.array([[26.830835356   , -6.960033590895],
                           [ 0.240134782174,  0.031972398386],
                           [ 2.971341010574, -2.141467087847],
                           [ 0.997482022276, -0.599895804555],
                           [ 0.470445852682, -0.482343299048]])
assert np.allclose(rating_matrix_d2[:5], d2_groundtruth)

rating_matrix_d5 = svd_transformed_rating(rating_df, 5)
d5_groundtruth = np.array([[26.828890919511, -6.950258673617,  1.515366172479, -3.295739836372,  4.157186247066],
                           [ 0.24012936551 ,  0.031884190966, -0.018132185754, -0.166903622477, -0.028071803934],
                           [ 2.9713046863  , -2.141818808721, -1.232252381363, -0.691324426529,  0.11961284882 ],
                           [ 0.997484724453, -0.59970688578 , -0.107441873278,  0.047329773205,  0.160751420569],
                           [ 0.470427947552, -0.482204394535,  0.062480138124, -0.238280610506,  0.055425249737]])
assert np.allclose(rating_matrix_d5[:5], d5_groundtruth)

rating_matrix_transformed = svd_transformed_rating(rating_df, 100)
assert rating_matrix_transformed.shape == (2770,100), "The transformed matrix is of the wrong dimension."
for i in range(100):
    for j in range(i):
        assert abs(rating_matrix_transformed.T[i].dot(rating_matrix_transformed.T[j])) < 1e-8, \
            "The column vectors in X* should be orthogonal."

```

A quick note on the above test cell: Ideally, the first $k'$ columns should stay the same for any $k$ as long as $k > k'$. However, you may notice small differences in the test cell (e.g. 26.830835356 vs. 26.828890919511). This is because the `TruncatedSVD` API uses a randomized SVD solver to speed up the calculation, so the output may lose minor precision.

As you can see, we even increased $k$ to 100, which is much larger than $k=2$ used in the document-word example. However, this is still much smaller than the 11,937 dimensions in the original rating matrix, and we shall soon see that these 100 dimensions have preserved most of the information in the original matrix.

### Exercise 6. (10 pts).
This exercise is to help you examine the power of SVD. That is, we want to see how the 100-dimension vectors preserve much information from the 11,937-dimensional vectors. Indeed, we shall see that the similarity between vectors is preserved after the SVD transformation. To show that, we are going to combine several techniques we have learned so far, including vector similarity, patterns in matrix data, and itemset similarity.

Please complete the following task:
1. Complete the `find_max_dot_prod_restaurants` and the `find_max_dot_prod_restaurants_transformed` function. Each function returns the `business_id` of the top-$n$ restaurants that have the largest dot product with Modavie. The `find_max_dot_prod_restaurants` function calculates the dot product based on the raw vectors (11,937 dimensions). The  `find_max_dot_prod_restaurants_transformed` function calculates with the transformed vectors (100 dimensions). You may copy and paste `find_max_dot_prod_restaurants` from Part III of this assignment.
2. Complete the `jaccard_similarity_before_after_svd_transform` function, which calculates the Jaccard similarity between the top-$n$ most similar restaurants (of course, this is an itemset!) before and after the SVD transformation.

We have provided all the test cases to help you verify your results.

```python
# This cell helps you prepare a new rating dataframe using the transformed vectors.
rating_matrix_transformed = svd_transformed_rating(rating_df, 100)
rating_df_transformed = pd.DataFrame(rating_matrix_transformed, index=rating_df.index)
rating_df_transformed.head()
```

<!-- grade_id: cell-869be40842fc6655, points: 0, grade: false, solution: true -->
```python
def find_max_dot_prod_restaurants(top_n):
    modavie_id = business_df[business_df.name.str.contains("Modavie")].index[0]
    modavie_vector = rating_df.loc[modavie_id]
    # YOUR CODE HERE
    raise NotImplementedError()

def find_max_dot_prod_restaurants_transformed(top_n):
    modavie_id = business_df[business_df.name.str.contains("Modavie")].index[0]
    modavie_vector_transformed = rating_df_transformed.loc[modavie_id]
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# This code block help you verify if `find_max_dot_prod_restaurants_transformed` is implemented correctly.
answer = find_max_dot_prod_restaurants_transformed(10)
assert answer.iloc[0]['name'] == "Schwartz's"
assert answer.iloc[1]['name'] == "La Banquise"
assert answer.iloc[2]['name'] == "Olive & Gourmando"
assert answer.iloc[3]['name'] == "Maison Christian Faure"
assert answer.iloc[4]['name'] == "Reuben's Deli & Steakhouse"

assert answer.iloc[5]['name'] == "Kazu"
assert answer.iloc[6]['name'] == "Belon"
assert answer.iloc[7]['name'] == "Au Pied de Cochon"
assert answer.iloc[8]['name'] == "L'Avenue"
assert answer.iloc[9]['name'] == "Poutineville"
```

<!-- grade_id: cell-e0ba9761d388a56b, points: 0, grade: false, solution: true -->
```python
def jaccard_similarity_before_after_svd_transform(top_n):
    max_sim_restaurants = find_max_dot_prod_restaurants(top_n)
    max_sim_restaurants_transformed = find_max_dot_prod_restaurants_transformed(top_n)
    # construct the set of the names of the most similar restaurants
    business_id_before = set(max_sim_restaurants.name)
    business_id_after = set(max_sim_restaurants_transformed.name)
    # compute the Jaccard similarity between the two set and return the Jaccard similarity
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return jaccard_similarity
```

<!-- grade_id: cell-cbb0be3036729dff, points: 10, grade: true, solution: false -->
```python
# This code block checks `find_max_dot_prod_restaurants`, `find_max_dot_prod_restaurants_transformed`, 
# and 'jaccard_similarity_before_after_svd_transform'. You should get the correct answer if all three 
# functions are implemented correctly.
# We hide some tests, so passing all the displayed assertions does not guarantee full points.

assert abs(jaccard_similarity_before_after_svd_transform(5) - 1) < 1e-8
assert abs(jaccard_similarity_before_after_svd_transform(10) - 0.8181818181818182) < 1e-8
assert abs(jaccard_similarity_before_after_svd_transform(2770) - 1) < 1e-8

```

As you can see, you can perform many tasks on the transformed vectors just as what you can do on the raw vectors. Since the transformed vectors have far fewer dimensions, the efficiency is much higher with the transformed vectors. Beyond efficiency, dimension reduction has many other benefits for dealing with matrix data.  You will learn these in the downstream machine learning classes, but let's stop here for now.  Matrix analysis has lots of applications in recommender systems.  We will revisit some of these techniques when that topic comes up.  

This concludes this assignment.